## CS3100 Assignment 1

* Released on Aug 27, 2021, 9pm
* Due on Sep 4, 2021, 11:59pm (moodle time).

### Important notes about grading:
1. **Compiler errors:** All code you submit must compile. Programs that do not compile will probably receive an automatic zero. If you are having trouble getting your assignment to compile, please visit the TA or instructor during the contact hours. If you run out of time, it is better to comment out the parts that do not compile, than submit a more complete file that does not compile.
2. **Late assignments:** Please carefully review the message sent on course groups about late submission so fassignments. All assignments handed in after the deadline will be considered late. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and realizing that you have done so after the deadline will be counted as a late submission.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```OCaml
let name = "Sathvick Reddy Beeravalli"
let rollno = "CS19B010"
```

Output:
```
val name : string = "Sathvick Reddy Beeravalli"

val rollno : string = "CS19B010"
```

## Problem 1

Write a function 

```ocaml
cond_dup : 'a list -> ('a -> bool) -> 'a list
```

list that takes in a list and preciate and duplicates all elements which satisfy the condition expressed in the predicate. For example, `cond_dup [3;4;5] (fun x -> x mod 2 = 1) = [3;3;4;5;5]`.

<!-- grade_id: cell-7950f1ea0a84b889, points: 0, grade: false, solution: true -->
```OCaml

let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)


let rec cond_dup l f lis =
  match l with
  |[] -> rev_list lis []
  |  h::t->
  if f h then cond_dup t f (h::h::lis) 
  else 
  cond_dup t f (h::lis)
   
let cond_dup l f =cond_dup l f []
  
  
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val cond_dup : 'a list -> ('a -> bool) -> 'a list -> 'a list = <fun>

val cond_dup : 'a list -> ('a -> bool) -> 'a list = <fun>
```

<!-- grade_id: cell-1d772b02ff64382c, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (cond_dup [3;4;5] (fun x -> x mod 2 = 1) = [3;3;4;5;5])
(* BEGIN HIDDEN TESTS *)
; assert (cond_dup [] (fun x -> x mod 2 = 1) = [])
; assert (cond_dup [1;2;3;4;5] (fun x -> x mod 2 = 0) = [1;2;2;3;4;4;5])
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 2

Write a function 

```ocaml
n_times : ('a -> 'a) * int * 'a -> 'a
```

such that `n_times (f,n,v)` applies `f` to `v` `n` times. For example, `n_times((fun x-> x+1), 50, 0) = 50`. If `n<=0` return `v`.


<!-- grade_id: cell-c28f98bd63f56ba1, points: 0, grade: false, solution: true -->
```OCaml
let rec n_times (f, n, v) =
  if n=0 then v
  else n_times (f,n-1,(f v))
  
  let n_times (f,n,v)=n_times(f,n,0)
```

Output:
```
val n_times : ('a -> 'a) * int * 'a -> 'a = <fun>

val n_times : (int -> int) * int * 'a -> int = <fun>
```

<!-- grade_id: cell-98364653c4b77b2d, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (n_times((fun x-> x+1), 50, 0) = 50)
(* BEGIN HIDDEN TESTS *)
; assert (n_times ((fun x->x+1), 0, 1) = 1)
; assert (n_times((fun x-> x+2), 50, 0) = 100)
(* END HIDDEN TESTS *)
```

Output:
```
Exception: Assert_failure ("[5]", 4, 2).
Raised at file "[5]", line 4, characters 2-43
Called from file "toplevel/toploop.ml", line 212, characters 17-27
```

## Problem 3

Write a function 

```ocaml
range : int -> int -> int list 
```

such that `range num1 num2` returns an ordered list of all integers from `num1` to `num2` inclusive. For example, `range 2 5 = [2;3;4;5]`. Raise the exception `IncorrectRange` if `num2 < num1`.

<!-- grade_id: cell-1fd147f7eb3b139d, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

exception IncorrectRange
 
let rec range num1 num2 lis=
  if num2<num1 then raise IncorrectRange
  else
    if num1=num2 then rev_list (num1::lis) []
    
    else range (num1+1) num2 (num1::lis)
    
    let range num1 num2 =range num1 num2 []
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

exception IncorrectRange

val range : int -> int -> int list -> int list = <fun>

val range : int -> int -> int list = <fun>
```

<!-- grade_id: cell-bba984120a807d42, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (range 2 5 = [2;3;4;5])
(* BEGIN HIDDEN TESTS *)
; assert (range 0 0 = [0])
; assert (try range 1 0 = [] with IncorrectRange -> true)
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 4

Write a function 

```ocaml
zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list
```

such that `zipwith f l1 l2` generates a list whose `ith` element is obtained by applying `f` to the `ith` element of `l1` and the `ith` element of `l2` . If the lists have different lengths, the extra elements in the longer list should be ignored. For example, `zipwith (+) [1;2;3] [4;5] = [5;7]`.

<!-- grade_id: cell-932054b41dace011, points: 0, grade: false, solution: true -->
```OCaml
let rec length l =
  match l with
  | [] -> 0
  | h::t -> 1 + length t

let rec zipwith f l1 l2 =
match l1 with 
|[]->[] 
|hd1::tl1->
match l2 with
 |[]->[]
 |hd2::tl2-> (f hd1 hd2)::zipwith f tl1 tl2 
 
```

Output:
```
val length : 'a list -> int = <fun>

val zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list = <fun>
```

<!-- grade_id: cell-55aefe5fa4f927d0, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (zipwith (+) [1;2;3] [4;5] = [5;7]) 
(* BEGIN HIDDEN TESTS *)
; assert (zipwith (fun x y -> (x,y)) [1;2;3;4] [5;6;7] = [(1,5); (2,6); (3,7)])
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 5

Write a function 

```ocaml
buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list 
```

that partitions a list into equivalence classes. That is, buckets equiv lst should return a list of lists where each sublist in the result contains equivalent elements, where two elements are considered equivalent if equiv returns true. For example:


```ocaml
buckets (=) [1;2;3;4] = [[1];[2];[3];[4]]
buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]]
buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]]
```

The order of the buckets must reflect the order in which the elements appear in the original list. For example, the output of `buckets (=) [1;2;3;4]` should be `[[1];[2];[3];[4]]` and not `[[2];[1];[3];[4]]` or any other permutation.

The order of the elements in each bucket must reflect the order in which the elements appear in the original list. For example, the output of `buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6]` should be `[[1;4];[2;5];[3;6]]` and not `[[4;1];[5;2];[3;6]]` or any other permutations.

Assume that the comparison function `('a -> 'a -> bool)` is commutative, associative and idempotent.

Just use lists. Do not use sets or hash tables.

List append function `@` may come in handy. `[1;2;3] @ [4;5;6] = [1;2;3;4;5;6]`.

<!-- grade_id: cell-f4d88f256cc9db8b, points: 0, grade: false, solution: true -->
```OCaml
let rec insert p ele lis temp check=
match lis with
|[]-> if check then temp else (temp@[[ele]])

|hd::tl-> match hd with 
          |[]->temp@[[ele]]
          |x::tl1 -> if p x ele then insert p ele tl (temp@[hd@[ele]]) true else insert p ele tl (temp@[hd]) check


let rec buc p l lis=
match l with
|[]->lis
|hd::tl->buc p tl (insert p hd lis  [] false) 
 
let buckets p l =buc p l [] 
```

Output:
```
val insert :
  ('a -> 'a -> bool) ->
  'a -> 'a list list -> 'a list list -> bool -> 'a list list = <fun>

val buc : ('a -> 'a -> bool) -> 'a list -> 'a list list -> 'a list list =
  <fun>

val buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list = <fun>
```

<!-- grade_id: cell-fdbd3b5564e151cb, points: 20, grade: true, solution: false -->
```OCaml
(* 20 points *)
assert (buckets (=) [1;2;3;4] = [[1];[2];[3];[4]])
; assert (buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]])
; assert (buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]])
(* BEGIN HIDDEN TESTS *)
; assert (buckets (fun x y -> (=) (x mod 2) (y mod 2)) [1;2;3;4;5;6] = [[1; 3; 5]; [2; 4; 6]])
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 6

Write a function 

```ocaml
remove_stutter : 'a list -> 'a list
```

that removes stuttering from the original list. For example, `remove_stutter [1;2;2;3;1;1;1;4;4;2;2]= [1;2;3;1;4;2]`. Option type may come in handy. 

<!-- grade_id: cell-8865e0dd05cbd32d, points: 0, grade: false, solution: true -->
```OCaml
let rec remove_stutter (l,a)=
match l with
|[]->[]
|hd::tl -> 
match a with Some(temp)->
if(hd=temp) then remove_stutter (tl,Some(hd))
else hd::remove_stutter (tl,Some(hd))
|None->hd::remove_stutter (tl,Some(hd))
let remove_stutter l =remove_stutter(l,None)
```

Output:
```
val remove_stutter : 'a list * 'a option -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-aa3fb845581c6fce, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (remove_stutter [1;2;2;3;1;1;1;4;4;2;2] = [1; 2; 3; 1; 4; 2]) 
(* BEGIN HIDDEN TESTS *)
; assert (remove_stutter [] = [])
; assert (remove_stutter [1;1;1;1;1] = [1])
; assert (remove_stutter [1;1;1;1;1;2] = [1;2])
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 7

Write a function

```ocaml
flatten : 'a list list -> 'a list
```

that flattens a list. For example, `flatten [[1;2];[3;4]] = [1;2;3;4]`.

<!-- grade_id: cell-1ba30bbe74fd3005, points: 0, grade: false, solution: true -->
```OCaml
let rec flatten l =
match l with 
|[]->[]
|hd::tl-> hd@flatten tl
```

Output:
```
val flatten : 'a list list -> 'a list = <fun>
```

<!-- grade_id: cell-8e986c3f5dbf5e89, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (flatten ([[1;2];[3;4]]) = [1;2;3;4])
(* BEGIN HIDDEN TESTS *)
; assert (flatten ([[1;2];[];[3;4];[]]) = [1;2;3;4])
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 8

Write a function

```ocaml

type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

fold_inorder : ('a -> 'b -> 'a) -> 'a -> 'b tree -> 'a
```

That does a inorder fold of the tree. For example,

`fold_inorder (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = [1;2;3]`

<!-- grade_id: cell-7170388040d2ef4b, points: 0, grade: false, solution: true -->
```OCaml
type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

let rec fold_inorder f acc t =
  match t with 
  |Leaf->acc
  |Node(x,y,z)->fold_inorder f(f (fold_inorder f acc x) y) z 
```

Output:
```
type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

val fold_inorder : ('a -> 'b -> 'a) -> 'a -> 'b tree -> 'a = <fun>
```

<!-- grade_id: cell-8bd22519c4af64bb, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (fold_inorder (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = [1;2;3]) 
(* BEGIN HIDDEN TESTS *)
; assert (fold_inorder (fun acc x -> x + acc) 0 (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = 6)
; assert (fold_inorder (fun acc x -> if acc < x then x else failwith "acc > x") 0 (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = 3)
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

## Problem 9

The usual recursive formulation of fibonacci function 

```ocaml
let rec fib n = 
  if n = 0 then 0
  else if n = 1 then 1
  else fib (n-1) + fib (n-2)
```

has exponential running time. It will take a long time to compute `fib 50`. You might have to interrupt the kernel if you did try to do `fib 50` in the notebook.

But we know that fibonacci number can be computed in linear time by remembering just the current `cur` and the previous `prev` fibonacci number. In this case, the next fibonacci number is computed as the sum of the current and the previous numbers. Then the program continues by setting `prev` to be `cur` and `cur` to be `cur + prev`.

Implement a tail recursive function `fib_tailrec` that uses this idea and computes the `nth` fibonacci number in linear time. 

```ocaml
fib_tailrec : int -> int
```

<!-- grade_id: cell-39286aaf6b703959, points: 0, grade: false, solution: true -->
```OCaml
let rec fib_tailrec (a,b,n) =
  if n==0 then a
  else
  if n==1 then b
  else fib_tailrec ((b),(a+b),(n-1))
 let fib_tailrec n=fib_tailrec(0,1,n)
```

Output:
```
val fib_tailrec : int * int * int -> int = <fun>

val fib_tailrec : int -> int = <fun>
```

<!-- grade_id: cell-57d3e7608b2286f5, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (fib_tailrec 50 = 12586269025)
(* BEGIN HIDDEN TESTS *)
; assert (fib_tailrec 100 = 3736710778780434371)
; assert (fib_tailrec 0 = 0)
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```